```python
# Initialize Otter
import otter
grader = otter.Notebook("DS3000-9000_A01.ipynb")
```

# Assignment 1: NumPy and Pandas

**DS 3000A/9000A Fall 2026**

This assignment will be graded on Gradescope.

1. Run `pip install otter-grader` first, if you don't have it installed.
2. Included in this notebook are public test cases, which you can use to verify your answers before submitting.
2. Additional private test cases will run on Gradescope upon submission. You will see your score but not the tests themselves.
3. You can revise and re-upload your submission on Gradescope up until the deadline.

## Toolbox

```python
## Install packages if neccessary.
## Uncomment and run if you need to.

# !pip install numpy pandas matplotlib
```

```python
import numpy as np
import pandas as pd
```

---

## Question 1

**25 Points**

Given two vectors $\mathbf{x}, \mathbf{y} \in \mathbb{R}^N$, their **_cosine similarity_** is defined by

$$
  \mathrm{sim}(\mathbf{x}, \mathbf{y}) = \frac{\mathbf{x}^\top \mathbf{y}}{\lVert\mathbf{x}\rVert \lVert\mathbf{y}\rVert}
$$

Write a Python function that computes the cosine similarity of its NumPy arrays arguments.

```python
def cosine_sim(x: np.ndarray, y: np.ndarray) -> float:
    """Cosine similarity.

    Args:
        x: A NumPy array of shape (N,).
        y: A NumPy array of shape (N,).

    Returns:
        The cosine similarity of x and y.
    """
    num = np.dot(np.transpose(x), y)
    den = np.linalg.norm(x) * np.linalg.norm(y)
    return num / den
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

---

## Question 2

**25 Points**

Given a matrix $\mathbf{A} \in \mathbb{R}^{M \times N}$, the **Gram matrix** is defined as

$$
\mathbf{G} = \mathbf{A}^\top \mathbf{A}
$$

The Gram matrix is always **symmetric** and **positive semi-definite**.

Write a function that computes and returns $\mathbf{G}$ given $\mathbf{A}$.
Then, write a function that returns the **diagonal elements** of $\mathbf{G}$ as a 1D NumPy array.

```python
def gram_matrix(A: np.ndarray) -> np.ndarray:
    """Compute the Gram matrix of A.

    Args:
        A: A NumPy array of shape (M, N).

    Returns:
        The Gram matrix of shape (N, N).
    """
    return np.dot(np.matrix_transpose(A), A)

def gram_diagonal(A: np.ndarray) -> np.ndarray:
    """Return the diagonal of the Gram matrix of A.

    Args:
        A: A NumPy array of shape (M, N).

    Returns:
        A 1D NumPy array of length N.
    """
    M = gram_matrix(A)
    return np.diag(M)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

---

## Question 3

**25 Points**

Using the provided `data.csv` file, complete the following:

1. Read the CSV into a DataFrame `df`. Note that the `apparentTemperature` column is in Fahrenheit.
2. Add a new column to `df`, called `temp_celsius` that contains `apparentTemperature` in Celsius using the conversion function:
$$
T_C = \frac{T_F - 32}{1.8}
$$
3. Assign to `df_hot_days` a filtered DataFrame containing only rows where `temp_celsius` **exceeds** $20°C$.
4. Assign to `mean_hot_temp` the **mean** of `temp_celsius` in `df_hot_days` as a Python `float`.

```python
df = pd.read_csv("data.csv")

df['temp_celsius'] = (df.apparentTemperature - 32) / 1.8

df_hot_days = df.copy()
df_hot_days = df_hot_days[df_hot_days.temp_celsius > 20]
mean_hot_temp = df_hot_days.temp_celsius.mean()
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

---

## Question 4
 
**25 Points**
 
Real-world data is messy. Two of the most common issues are **missing values** and **categorical variables**.
 
We provide a small dataset:

```python
df_raw = pd.DataFrame(
    {
        "temperature": [22.1, np.nan, 18.4, np.nan, 25.0],
        "weather": ["sunny", "cloudy", "rainy", "cloudy", "sunny"],
        "humidity": [0.4, 0.6, np.nan, 0.7, 0.5],
    }
)
```

Complete the following:
 
1. Assign to `n_missing` the **total number of missing values** in `df_raw` as a Python `int`.
2. Assign to `df_clean` a copy of `df_raw` where missing values in `temperature` and `humidity` are **filled with the column mean**.
3. Assign to `df_encoded` the `df_clean` DataFrame with the `weather` column replaced by **one-hot encoding** (use `pd.get_dummies`). Drop the original `weather` column.

```python
n_missing = int(df_raw.isnull().sum().sum())

df_clean = df_raw.copy()
df_clean['humidity'] = df_clean['humidity'].fillna(df_clean['humidity'].mean())
df_clean['temperature'] = df_clean['temperature'].fillna(df_clean['temperature'].mean())

df_encoded = pd.get_dummies(df_clean, columns=['weather'])
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!

    q2 results: All test cases passed!

    q3 results: All test cases passed!

    q4 results: All test cases passed!
<IPython.core.display.HTML object>
```

Tests for q1 (points: 25)
```
>>> x = np.array([1, 2, 3])
>>> y = np.array([4, 5, 6])
>>> assert np.isclose(cosine_sim(x, y), 0.9746318461970762)
```
```
>>> x = np.array([1, 0, 0])
>>> y = np.array([0, 1, 0])
>>> assert np.isclose(cosine_sim(x, y), 0.0)
```


Tests for q2 (points: 25)
```
>>> A = np.array([[1, 2, 3], [4, 5, 6]], dtype=np.float64)
>>> G = gram_matrix(A)
>>> assert G.shape == (3, 3)
```
```
>>> A = np.array([[1, 2, 3], [4, 5, 6]], dtype=np.float64)
>>> d = gram_diagonal(A)
>>> assert d.shape == (3,)
```


Tests for q3 (points: 25)
```
>>> assert 'temp_celsius' in df.columns
```
```
>>> assert (df_hot_days['temp_celsius'] > 20).all()
```
```
>>> assert isinstance(mean_hot_temp, float)
```


Tests for q4 (points: 25)
```
>>> assert isinstance(n_missing, int)
>>> assert n_missing == 3
```
```
>>> assert df_clean.isnull().sum().sum() == 0
```
```
>>> assert np.isclose(df_clean['temperature'].iloc[1], df_raw['temperature'].mean(skipna=True))
>>> assert np.isclose(df_clean['humidity'].iloc[2], df_raw['humidity'].mean(skipna=True))
```
```
>>> assert 'weather' not in df_encoded.columns
```
```
>>> assert 'weather_sunny' in df_encoded.columns
>>> assert 'weather_cloudy' in df_encoded.columns
>>> assert 'weather_rainy' in df_encoded.columns
```
